# MiniProyecto 2: Segmentación de higado en volumenes MRI

## Entrega 2:
En la anterior entrega hemos desarrollado un algoritmo completo de segmentación para el hígado en MRIs. Los resultados no son perfectos y tendremos numerosos falsos positivos, falsos negativos, huecos, protuberancias y bahias dentro de nuetras mascaras de segmentación que disminuirán nuestras metricas de evaluación. Para poder refinar nuestros resultados tenemos 2 posibilidades, mejorar las imagenes originales para facilitar el trabajo del algoritmo de predicción en las zonas de peor rendimiento o mejorar la mascara de predicción con morfología matematica para pulir las imperfecciones y subir las métricas de evaluación.

### Parte 1: Operación morfológica como pre-procesamiento
Hasta ahora, hemos realizado un proceso de pre-procesamiento basado unicamente en la corrección de intensidades y la ecualización o especificación de histogramas. Estos procesos se basan casi exclusivamente en las intensidades de los pixeles, obviando la posible relación entre los pixeles y sus vecindades. En esta ocasión, usaremos las herramientas de la morfología matematica pero aplicada a **escala de grises** como metodo de pre-procesamiento.

#### Parte 1.1: Visualización de errores
Antes de realizar cualquier pre-procesamiento es importante entender nuestras imagenes y las limitaciones de nuestro modelo para poder diseñar una metodología eficaz que mejore nuestros resultados y no nos resulte perjudicial. Tampoco nos funcionará realizar una metodología de pre-procesamiento a ciegas, siempre es importante comprender nuestro modelo y nuestras imagenes para no solo diseñar un metodo que nos funcione en una base de datos en especifico, sino que sea los mas generalizable posible. 


Cargue el volumen segmentation-1.nii.gz y el volumen volume-1.nii.gz de la base de datos de entrenamiento, luego utilice la función de umbralización creada en la entrega anterior utilizando un valor minimo de 30 y maximo de 220 para generar una predicción. Por ultimo, visualize las estructuras de el **volumen original** que usted esta segmentando con su algoritmo. Haga una función que cree 2 volumenes adicionales, uno que contenga unicamente los pixeles falsos positivos y otro que contenga unicamente los pixeles falsos negativos. Para el subplot encuentre en cada volumen los slices (axial, transversal y sagital) con mayor cantidad de errores.

> **Nota:** Recuerde que los FP serán aquellos pixeles que se encuentran en su prediccion pero no en su groundtruth. Los FN serán los pixeles que se encuentran en su groundtruth pero no en su predicción.

```python
import os
import numpy as np
from glob import glob
import nibabel as nib
from matplotlib import pyplot as plt

from tqdm import tqdm

from skimage.measure import label
from skimage.morphology import ball, erosion, dilation, closing, opening, white_tophat, black_tophat
```

<!-- grade_id: cell-f20da4e0ea392d3d, points: 0, grade: false, solution: true -->
```python
volume1='' #volumen 1 de entrenamiento
segmen1='' #segmentacion del volumen 1 de entrenamiento, groundtruth
segment_pred='' #volumen de prediccion de su algoritmo
FN_volumes='' #volumen unicamente con pixeles Falsos Negativos
FP_volumes='' #volumen unicamente con pixeles Falsos Positivos

def errores(segment_pred,segmen1):
    'segment_pred (ndarray): volumen de predicción'
    'segmen1 (ndarray): volumen de groundtruth'
    # YOUR CODE HERE
    raise NotImplementedError()
    return FP_volumes,FN_volumes
    
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6eecdf09f1acaf4c, points: 0.5, grade: true, solution: false -->
```python
assert np.isclose(np.sum(volume1),-1550711485.0,0.01), 'No se abrió el volumen 1'
assert len(np.unique(segmen1))==2, f'Las etiquetas del volumen de anotacion deben ser 2 y son {len(np.unique(segmen1))}, puede que haya guardado el volumen 1 en esta variable o no haya contado la etiqueta 1 y 2 como 1.'
aux_segment=vol_max_component(volume1,220,0)
fp,fn=errores(aux_segment,segmen1)
assert np.isclose(np.sum(fp),265198), f'La cantidad de FP deberian ser 326707 y son {np.sum(fp)}'
assert np.isclose(np.sum(fn),1333), f'La cantidad de FN deberian ser 481 y son {np.sum(fn)}'
```

#### Parte 1.2: Operaciones morfologicas como pre-procesamiento
Hemos podido observar aquellos pixeles que son errores en nuestra predicción, ahora debemos realizar un pre-procesamiento en escala de grises sobre los volumenes originales para facilitar el trabajo de segmentación de nuestro modelo. En la siguiente celda realize el tophat con un elemento estructurante de *ball* de radio 3. Vea https://scikit-image.org/docs/stable/api/skimage.morphology.html para más información. Visualize el resultado en sus 3 ejes en los slices que desee.

<!-- grade_id: cell-02e14924565a38b8, points: 0, grade: false, solution: true -->
```python
vol_pre='' #volumen pre-procesado 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ef67a7a810a2a9a4, points: 0.3, grade: true, solution: false -->
```python
assert len(np.unique(vol_pre))>2, 'Deberia estar haciendo el pre-procesamiento en el volumen original, no en el predicho ni en el groundtruth'
assert np.isclose(np.sum(vol_pre!=0),2114154,0.001), 'No calculó bien el pre-procesamiento'
```

Podemos observar que el resultado anterior son las estructuras más brillantes que son de menor tamaño a nuestro elemento estructurante. Ahora que los tenemos extraidos, umbralice esta mascara para obtener solo los valores entre 70 y 1024, y elimine estas estructuras del volumen original para no tenerlas en cuenta.

> **Sugerencia:** Usted quiere eliminar los pixeles del volumen original que se encuentran en la nueva mascara creada, por lo que el objetivo será restar a cada pixel su **misma intensidad**. Piense como obtener la intensidad original de los pixeles de la mascara. 

Compare el resultado de la segmentación que usted propuso en un inicio con el volumen orignial y con este nuevo volumen. Haga un subplot donde se pueda ver el resultado en los 3 ejes.

<!-- grade_id: cell-5f693f90b65c546e, points: 0, grade: false, solution: true -->
```python
vol_pre_total='' #volumen pre-procesado completamente
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-502b51eb7f87fd2d, points: 0.4, grade: true, solution: false -->
```python
assert len(np.unique(vol_pre_total))!=2, 'No hizo el pre-procesamiento completo en el volumen original sino en otro volumen'
assert np.min(vol_pre_total)>=-1024, 'Seguramente restó la mascara binaria directamente'
assert np.isclose(np.sum(vol_pre_total>0.8),278227,0.001), 'No realizó bien el pre-procesamiento para eliminar las estructuras en el volumen original'
```

Veamos como afecta este preprocesamiento a nuestro metodo de segmentación.

<!-- grade_id: cell-b6baefea6922f13b, points: 0, grade: false, solution: true -->
```python
def Jaccard_3D(Im,Gt):
    '''
    :param Im: Volumen de prediccion.
    :param Gt: Segmentacion groundtruth.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return J
```

```python

J_sinpre=Jaccard_3D(aux_segment,segmen1)
segment_conpre=vol_max_component(vol_pre_total,200,30)
J_conpre=Jaccard_3D(segment_conpre,segmen1)

print(f'El jaccard sin preprocesamiento es: {J_sinpre}')
print(f'El jaccard con preprocesamiento es: {J_conpre}')
```

```python
fp2,fn2=errores(segment_conpre,segmen1)
print(f'Los falsos positivos del volumen no preprocesado es {np.sum(fp)}')
print(f'Los falsos positivos del volumen preprocesado es {np.sum(fp2)}')

print(f'Los falsos negativos del volumen no preprocesado es {np.sum(fn)}')
print(f'Los falsos negativos del volumen preprocesado es {np.sum(fn2)}')
```

Como podemos observar, el preprocesamiento aumentó los falsos negativos pero disminuyó los falsos positivos. El procesamiento que realizamos se ve beneficiado si se sobresegmenta.

El anterior procesamiento fue tan solo un ejemplo de lo que podemos lograr hacer con la morfología matematica en escala de grises. Planee una linea de base similar que le pueda ayudar a eliminar elementos que no le interesan o resaltar los objetos que si. Existen muchos algoritmos de morfología matematica y la combinación de ellos nos permiten hacer procesamientos verdaderamente interesantes. Luego, haga experimentos (minimo 4) con el tamaño, la forma y las dimensiones (2D o 3D) del elemento estructurante. El elemento estructurante que nosotros usamos era una Bola de radio 3, un elemento de 3 dimensiones que tiene skimage.morphology, ustedes pueden usar estos elementos estrcutrantes o diseñar uno como una matriz de numpy.

>**Nota:** Cree una linea de base verdaderamente util y que le ayude a mejorar sus metricas, no haga un procesamiento solo por hacerlo, de igual forma para los experimentos con el elemento estructurante. Observe sus volumenes de FP y de FN, seguramente les dará una idea.


>**Recuerde:** Estos procesamientos se hacen como pre-procesamiento, es decir, en el volumen original y no en la segmentación ya hecha.

<!-- grade_id: cell-5d65f109c045f18f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Por ultimo, cree una función que resuma por completo su preprocesamiento. Recuerde que la idea es que se mejore nuestra segmentación final, esta será parte de su nota. Cree la función que entre por parametro un volumen unicamente. El resto de parametros como elementos estructurantes, tamaños, umbrales, etc. deben ir por defecto.

<!-- grade_id: cell-ea5b21a9f7c6f20a, points: 0, grade: false, solution: true -->
```python
def preprocesamiento_total(vol):
    # YOUR CODE HERE
    raise NotImplementedError()
    return volume_preprocessed
```

```python
mean_jaccard = 0.0
vol_val_paths = glob(os.path.join('DB','valid','volume_*.nii.gz'))
seg_val_paths = glob(os.path.join('DB','valid','segmentation_*.nii.gz'))
pbar = tqdm(total=len(vol_val_paths))
for i, (vol_val_path, seg_val_path) in enumerate(zip(vol_val_paths, seg_val_paths), start = 1):
    vol_val = nib.load(vol_val_path).get_fdata()
    seg_val = nib.load(seg_val_path).get_fdata()>0
    vol_pre_total2 = preprocesamiento_total(vol_val)
    vol_segmentado = vol_max_component(vol_pre_total2, 220, 30)
    jaccard_valid = Jaccard_3D(vol_segmentado,seg_val)
    mean_jaccard += (jaccard_valid - mean_jaccard) / i
    pbar.update(1)
pbar.close()
print(f'El jaccard medio de validacion es: {mean_jaccard}')
```

<!-- grade_id: cell-0a1502af3c3c394c, points: 0.3, grade: true, solution: false -->
```python
assert mean_jaccard>0.37, 'El jaccard de su preprocesamiento total debe ser mayor al preprocesamiento propuesto'
```

### Parte 2: Operación morfológica como post-procesamiento
Otra metodología que podemos llevar a cabo para mejorar nuestras mascaras de segmentación es la morfología binaria en las mascaras de segmentación. A diferencia de el pre-procesamiento en escala de grises, que busca resaltar o eliminar elementos en la imagen original para facilitar el proceso de prediccion del modelo, el post-procesamiento busca eliminar imperfecciones de las mascaras predichas por el modelo. En ese sentido, el post-procesamiento lo diseñaremos en base a las imperfecciones de la mascara y conociendo la forma de nuestro objeto de interes, olvidandonos por completo de nuestra imagen original.

**Importante:** Para disminuir los tiempos de carga utilice la funciones especialmente diseñadas para imágenes o volúmenes binarios.

#### Parte 2.1: Visualización de la mascara binaria
Para poder diseñar una linea de base con la que podamos mejorar nuestra predicción final primero debemos visualizar los errores de nuestra mascara. Recuerden que, al igual que los filtros, un procesamiento puede ayudar a nuestro resultado final o empeorarlo, por lo que hacer un procesamiento a ciegas sin conocer las limitaciones de nuestro modelo es una mala idea. Recorra nuevamente todo su volumen de predicción con el mejor preprocesamiento encontrado en el punto anterior y obtenga un volumen que tenga los pixeles falsos positivos y los pixeles falsos negativos como en el primer punto. Visualice de la misma forma los slice con más FP y FN junto con la mascara original y su mascara predicha en un subplot.

```python
# importar funciones importantes
from skimage.morphology import binary_closing, binary_opening, binary_erosion, binary_dilation
```

<!-- grade_id: cell-6642352d900dfd9c, points: 0, grade: false, solution: true -->
```python
vol_preprocesado='' #volumen final pre-procesado
vol_segmentado='' #volumen segmentado con el pre-procesamiento
FP_sinpost='' #volumen con pixeles unicamente Falsos Positivos
FN_sinpost='' #volumen con pixeles unicamente Falsos Negativos

# YOUR CODE HERE
raise NotImplementedError()
```

#### Parte 2.2: Operaciones morfologicas como post-procesamiento
Ahora, ya usted fué capaz de observar cuales fueron sus principales fallos en su mascara de segmentación con un volumen pre-procesado. Dependiendo de su caso, use una apertura o una clausura con una bola de radio 1 para mejorar su resultado.

<!-- grade_id: cell-0c91c8aa63bc584a, points: 0, grade: false, solution: true -->
```python
vol_post='' #volumen post-procesado
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d333b3754fd1e8eb, points: 0.3, grade: true, solution: false -->
```python
assert len(np.unique(vol_post))==2, 'No esta post-procesando en el volumen binario predicho'
assert Jaccard_3D(vol_post,segmen1)>Jaccard_3D(vol_segmentado,segmen1), 'Su post-procesamiento no mejora su predicción'
```

Realice tambien 2 experimentos adicionales con las operaciones basicas, *erosion* y *dilatacion* para mejorar su mascara de segmentación. Los procesos que elija dependerán mucho de si usted sobresegmentó o subsegmentó. Experimente con tamaños y dimensiones del elemento estructurante o incluso, con aperturas y clausuras iterativas.

<!-- grade_id: cell-172baafb0aec0ead, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Un procedimiento muy interesante que podemos llevar a cado es la recontrucción geodesica para obtener unicamente el hígado. El procedimiento es el siguiente:

>1. Erosionamos nuestro volumen con un elemento estructurante grande. Use una *ball* de radio 5. Esto nos dará como resultado los pixeles de los elementos más grandes del volumen.
>2. Usaremos un elemento estructurante pequeño para realizar una dilatación del resultado anterior. Use una *ball* de radio 1. 
>3. A el resultado del punto **2** le realizaremos una intersección con la mascara original, el volumen previo a la erosión con el elmento grande.
>4. Repita los puntos del 2 al 3 de forma iterativa hasta que el resultado no cambie.

Use la siguiente celda para realizar la reconstrucción geodesica de su volumen luego de la apertura del punto anterior y calcule el Jaccard3D para cuantificar la mejora. Cree una función que entre como parametro la segmentación como mascara, el elemento estructurante E1 con el que se hará la erosión, el elemento estructurante E2 con el que se hará la dilatación y un numero maxximo de iteraciones para no tener loops infinitos. Debe retornar el resultado de la dilatación geodesica y el numero de iteraciones.

<!-- grade_id: cell-a6aeeb684c2c008e, points: 0, grade: false, solution: true -->
```python
def geodesic_dilatation(mask, E1,E2,max_iterations):
    # YOUR CODE HERE
    raise NotImplementedError()
    return result,iterations
```

<!-- grade_id: cell-a92ac7907fd27678, points: 0, grade: false, solution: true -->
```python
vol_post_2='' #volumen de postprocesamiento total previo a la dilatación geodesica
E1='' # elemento estructurante para la erosión
E2='' # elemento estructurante para la dilatación
max_iterations='' #iteraciones maximas
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9057c98f13c905d4, points: 0.1, grade: true, solution: false -->
```python
assert np.sum(E1)==515, 'La erosión se está realizando con un elemento erroneo'
assert np.sum(E2)==7, 'La dilatacíon se está realizando con un elemento erroneo'
```

```python
reconstruccion,iterations=geodesic_dilatation(vol_post,E1,E2,max_iterations) # Esta función puede demorar en correr alrededor de 5 minutos
```

```python
Jaccard_3D(reconstruccion,segmen1)
```

<!-- grade_id: cell-80e13ca3f374d82d, points: 0.4, grade: true, solution: false -->
```python
assert iterations<max_iterations, 'Sus maximas iteraciones deben ser mayores a las iteraciones completas'
assert np.sum(reconstruccion), 'Su volumen de reconstrucción está vacio'
assert np.sum(reconstruccion)!=np.sum(vol_post), 'No realizó ningun procesamiento'
assert np.isclose(np.sum(dilation(reconstruccion.astype(int), E2)*vol_post-reconstruccion.astype(int)),0) or iterations==max_iterations, f'La dilatación de la reconstrucción intesección con la mascara no es igual a la reconstrucción, no se realizó bien la geodesic_dilatation'
```

Nuestras metricas seguramente hayan mejorado con este post-procesamiento, ya que eliminamos todos los elementos menores a nuestro elemento estructurante sin dañar tanto nuesto organo de interes. Sin embargo, nuestra mascara aun está lejos de ser perfecta, diseñe una linea de base con la que podrá mejorar su mascara de segmentación usando los algoritmos de skimage.morphology. Luego, realice algunos experimentos (minimo 4) con los que podrá mejorar su resultado final, modifique el tamaño del elemento estructurante, la forma y las dimensiones (2D y 3D). Pruebe tambien otras operaciones morfologicas como el relleno de huecos o el miss or hit si es necesario. Recuerde que nuestro baseline de tomar el elemento conexo mayor es un post-procesamiento, modifiquelo como crea oportuno tambien. Parte de su nota será mejorar la segmentación por medio del post-procesamiento

<!-- grade_id: cell-3c14be816e0ad3e1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Por ultimo, realice una función que genere el postprocesamiento de un volumen binario que entre por parametro que realice todo el procedimiento que usted planteó.

<!-- grade_id: cell-714b32ef7d3e4b41, points: 0, grade: false, solution: true -->
```python
def postprocesamiento_total(vol):
    # YOUR CODE HERE
    raise NotImplementedError()
    return vol_post
```

```python
mean_jaccard = 0.0
vol_val_paths = glob(os.path.join('DB','valid','volume_*.nii.gz'))
seg_val_paths = glob(os.path.join('DB','valid','segmentation_*.nii.gz'))
pbar = tqdm(total=len(vol_val_paths))
for i, (vol_val_path, seg_val_path) in enumerate(zip(vol_val_paths, seg_val_paths), start = 1):
    vol_val = nib.load(vol_val_path).get_fdata()
    seg_val = nib.load(seg_val_path).get_fdata()>0
    vol_pre_total2 = preprocesamiento_total(vol_val)
    vol_segmentado = vol_max_component(vol_pre_total2, 200, 30)
    vol_post = postprocesamiento_total(vol_segmentado)
    jaccard_valid = Jaccard_3D(vol_post,seg_val)
    mean_jaccard += (jaccard_valid - mean_jaccard) / i
    pbar.update(1)
pbar.close()
print(f'El jaccard medio de validacion es: {mean_jaccard}')
```

<!-- grade_id: cell-3fd8f3f320a0a3d1, points: 0.2, grade: true, solution: false -->
```python
assert mean_jaccard>0.46, 'El Jaccard del preprocesamiento propuesot por usted debe ser mayor al propuesto por nosotros'
```

## Entregables
Los entregables para esta entrega son un Jupyter notebook (.ipynb), un informe (.pdf) y un archivo de texto (.txt).

El Jupyter notebook debe estar completamente resuelto. Recuerde que la nota depende del funcionamiento de su código. El archivo de texto corresponde a las secciones de código de su Jupyter notebook, para más información sobre cómo convertir las secciones de código del Jupyter notebook a archivo de texto ver el video en el siguiente [enlace](https://uniandes-edu-co.zoom.us/rec/share/QOxUUIw7Uz9DsnKPyftOXZgM4bx5d7KHMYK-dqvGnkyJ2GfqvT7NU3lhQc0NtNSU.jveWB2H34S4uLIoB).

En este miniproyecto debe hacer un informe con Introducción, Resultados y análisis, y Conclusiones respondiendo las preguntas en las secciones donde usted crea pertinente responderlas. Toda la información sobre secciones, preguntas y elementos adicionales está en la siguiente sección de informe. La extensión máxima del informe es de 4 páginas. El informe debe utilizar el formato que puede encontrar siguiente [enlace](https://github.com/cvpr-org/author-kit). Si tiene dudas de cómo utilizar el formato puede preguntar al Asistente Graduado.

## Informe
Ya hemos realizado todo un procedimiento de pre-procesamiento y post-procesamiento para mejorar los resultados finales de nuestro modelo. Recuerde que queremos observar la sensibilidad de nuestro modelo a los diferentes parametros que hemos probado para así obtener el mejor proceso posible y habrán ocasiones en las que combinaciones de parametros funcionarán mejor que los parametros individuales, tengalo en cuenta. 

No olvide contestar las siguientes preguntas en su informe:

### Operación morfológica como pre-procesamiento
* Teniendo en cuenta nuestro metodo de segmentación usando el elemento conexo más grande ¿De que forma es útil el preprocesamiento? ¿Como explica el aumento en el indice de Jaccard?
* ¿Por qué nuestro procesamiento (el realizado en la sección 1.2, antes de la función de preprocesamiento total) aumenta los falsos negativos pero disminuye los falsos positivos?
* ¿Cree que el pre-procesamiento con escala de grises es necesario? ¿Puede obtener resultados similares usando simplemente el post-procesamiento? 
* Describa sus experimentos y explique su razon de ser, analice sus resultados. Incluya tambien una explicación de su pre-procesamiento total.

### Operación morfológica como post-procesamiento
* ¿Que errores en la mascara binaria permanecieron aun con el uso del preprocesamiento?
* ¿El procedimiento propuesto de la reconstrucción geodesica funcionará para segmentar algun otro organo que no sea el higado? ¿Que limitaciones tiene esta metodología?
* ¿Sus resultados son mayores con pre-procesamiento, con post-procesamiento o con ambos? ¿Por que?
* Describa sus experimentos y explique su razon de ser, analice sus resultados. Incluya tambien una explicación de su post-procesamiento total.

```python
from utils import *
converter("Entrega 2")
```